# **Homework 2 Problem 1 (22 points) [TA: Anurata Hridi]**

In this problem, you'll be exploring decision trees. 

## Instructions:
1. Use github.ncsu.edu to submit your work repository if you have not done so yet. It is your responsbility to ensure the TAs have access to your work before the deadline. Make sure the repo is private.
2. Do not modify the code structure given. Answer the questions in the designated space.

All the best!

## 1.0) Loading Data and Libraries

```python
import numpy as np

import pandas as pd
import matplotlib.pyplot as plt
import sklearn
from sklearn.tree import DecisionTreeClassifier
from sklearn.tree import plot_tree
from sklearn.model_selection import train_test_split
# Remember you have to run this cell block before continuing!
```

Output:
```
執行具有 'Python 3.13.0' 的儲存格需要 ipykernel 套件。

執行下列命令以將 'ipykernel' 安裝到 Python 環境中。

命令: '/bin/python3.13 -m pip install ipykernel -U --user --force-reinstall'
```

## Important Note about Non-Determinism (Randomness)
Here on forward, we're going to be using a lot of methods that have some randomness in their output. To address this, most sklearn libraries/functions allow you to a `random_state` parameter, which allows the methods to have consistent output each time you run them. Make sure to set `random_state=random_seed` so we can reproduce and check your work.

```python
# set a seed for reproducibility
random_seed = 25
np.random.seed(random_seed)
```

## 1.1) Building Trees & Interpretation (2 pts)

**Your goal is to create a set of *human-interpretable* rules to tell seeds apart**. 

You will make a decision tree using the heart diseases dataset, evaluate its performance, and interpret its rules.

```python
# Load the data
# Read the seeds dataset and translate to pandas dataframe
from sklearn import datasets
from sklearn.preprocessing import MinMaxScaler

# Define column names for the dataset
column_names = ['Area', 'Perimeter', 'Compactness', 'KernelLength', 'KernelWidth', 'AsymmetryCoefficient', 'KernelGrooveLength', 'Class']

# Load the dataset from the URL
url = 'https://archive.ics.uci.edu/ml/machine-learning-databases/00236/seeds_dataset.txt'
seeds_df = pd.read_csv(url, delim_whitespace=True, header=None, names=column_names)

# Make sure data is in the same range
seeds_temp_df = MinMaxScaler().fit_transform(seeds_df.iloc[:, :-1])

seeds_temp_df = pd.DataFrame(seeds_temp_df, columns = ['Area', 'Perimeter', 'Compactness', 'KernelLength', 'KernelWidth', 'AsymmetryCoefficient', 'KernelGrooveLength'])
# Note that the "class" attribute is type of wheat
seeds_df = pd.concat([seeds_temp_df, seeds_df["Class"]], axis=1)
```

```python
seeds_df 
```

### Building the Tree
The fraction of test data should be 20%.

As always, use `random_state=random_seed` for all of your model seeds to make sure your ouput is reproducable.

```python
np.random.seed(random_seed)
test_data_fraction = 0.2
```

<!-- grade_id: train_test_split, points: 0, grade: false, solution: true -->
```python
# Goal: Split the dataset into features and class labels
# and split the X and Y datasets into training and testing sets
from sklearn.model_selection import train_test_split

# TODO: Make sure to define these variables in your solution below
X_train = X_test = Y_train = Y_test = None

### BEGIN SOLUTION

### END SOLUTION
```

<!-- grade_id: train_test_split-public, points: 2, grade: true, solution: false -->
```python
assert(np.shape(X_train) == (168, 7))
assert(len(X_test) == 42)
```

<!-- grade_id: train_test_split-private, points: 1, grade: true, solution: false -->
```python
# Note: We will check your code here against hidden test cases, too.  
```

Now use the diagram generated to interpret the tree. If you had to write some human-interpretable rules for classifying seeds, what would they be?

# 2) Pruning & Hyperparameters

## 2.1 Example: Pre-Pruning (4 pts)

Below we will give some examples of how to use pre-pruning hyperparameters. You will answer some written questions throught this section.

Below we have written some example code to train a decision tree, plot it, and calculate its accuracy.

**Note**: We use the parameter `max_depth=2` to constrain the size and complexity of your decision tree. Remember this is a form of *pre-pruning*.

**Hint**: You may need to right-click and download the figure in order to view it. One some versions, you may need to shift-right-click.


```python
# Goal: Train a decision tree, plot it, and calculate its accuracy.
from sklearn.metrics import accuracy_score
gini_tree = DecisionTreeClassifier(criterion = "gini", max_depth=2, random_state=random_seed).fit(X=X_train, y=Y_train)
plt.figure(figsize=(40,15))
seeds_features = seeds_df.iloc[:,0:-1]
_ = plot_tree(gini_tree, feature_names=seeds_features.columns)
print(f'Train Accuracy: {accuracy_score(Y_train, gini_tree.predict(X_train))}')
print(f'Test Accuracy: {accuracy_score(Y_test, gini_tree.predict(X_test))}')
```

In the above example, we set the `max_depth=2`. Play around with different values for `max_depth`. How does the training/test accuracy change? What is the biggest the tree can be (without pre-pruning)? What do you think is the right size of the tree? Why? **Discuss and explain your reasoning here.**

**Answer here.**

### Min Impurity Decrease

Read the documentation for [DecisionTreeClassifier](https://scikit-learn.org/stable/modules/generated/sklearn.tree.DecisionTreeClassifier.html) and take note of the various params that can be used for pre-pruning. We will go in detail on `min_impurity_decrease` together.

This parameter set the minimium amount that the impurity (e.g. GINI Index) must go down when we split on a node, for that split to be allowed. If the minimum is 0 (default), this isn't a restriction. If it is high (e.g. 0.3), only very effective splits will be allowed, and we'll have a smaller tree. Note that this is still *pre-pruning* because we stop the tree from growing.

```python
impurity_decrease = 0.05
gini_tree = DecisionTreeClassifier(criterion = "gini", random_state=random_seed, min_impurity_decrease=impurity_decrease).fit(X=X_train, y=Y_train)
plt.figure(figsize=(20,10))
_ = plot_tree(gini_tree, feature_names=seeds_features.columns)
print(f"Accuracy with Impurity Decrease {impurity_decrease}")
print(f'Train Accuracy: {accuracy_score(Y_train, gini_tree.predict(X_train))}')
print(f'Test Accuracy: {accuracy_score(Y_test, gini_tree.predict(X_test))}')
```

```python
# However, if you go too far you end up decreasing training accuracy
impurity_decrease = 0.1
gini_tree = DecisionTreeClassifier(criterion = "gini", random_state=random_seed, min_impurity_decrease=impurity_decrease).fit(X=X_train, y=Y_train)
plt.figure(figsize=(20,10))
_ = plot_tree(gini_tree, feature_names=seeds_features.columns)
print(f"Accuracy with Impurity Decrease {impurity_decrease}")
print(f'Train Accuracy: {accuracy_score(Y_train, gini_tree.predict(X_train))}')
print(f'Test Accuracy: {accuracy_score(Y_test, gini_tree.predict(X_test))}')
```

There are other methods for pre-pruning, but we won't explore them today.
Read the documentation for [DecisionTreeClassifier](https://scikit-learn.org/stable/modules/generated/sklearn.tree.DecisionTreeClassifier.html) and take note of the various params that can be used for pre-pruning. These are

* `max_depth`
* `min_samples_split`
* `min_samples_leaf`
* `max_features`
* `max_leaf_nodes`
* `min_impurity_decrease` (already covered)

## 2.2 Post Pruning [Cost Complexity Pruning] (3 pts)

Sklearn implements a more advanced form of post-pruning that we haven't covered in class called [cost complexity pruning (CCP)](https://scikit-learn.org/stable/auto_examples/tree/plot_cost_complexity_pruning.html). The idea is essentially a more complex/robust form of what we've talked about in the seminars - seeing if estimated error improves after pruning certain nodes after the tree is built. The amount of the tree pruned is controlled by the paramater `ccp_alpha`. 

Below, complete the `dt_performace` function to train and evaluate a decision tree with the given `ccp_alpha` value. Then try some different values of `ccp_alpha` in the range (0,0.3) to see their effect on the accuracy.

<!-- grade_id: dt_performance, points: 0, grade: false, solution: true -->
```python
def dt_performance(ccp_alpha, X_train, Y_train, X_test, Y_test, random_seed):
    gini_tree = train_acc_value = test_acc_value = None
    # Build a decision tree using ccp_alpha
    ### BEGIN SOLUTION

    ### END SOLUTION
    return gini_tree, train_acc_value, test_acc_value


ccp_alpha = 0.1
gini_tree, trainin_acc_value, test_acc_value = dt_performance(ccp_alpha, X_train, Y_train, X_test, Y_test, random_seed)

# Now test with multiple ccp_alpha values and see how training and testing accuracy change.
print(f"Accuracy with ccp_alpha={ccp_alpha}")
print(f'Train Accuracy: {accuracy_score(Y_train, gini_tree.predict(X_train))}')
print(f'Test Accuracy: {accuracy_score(Y_test, gini_tree.predict(X_test))}')

plt.figure(figsize=(20,10))
_ = plot_tree(gini_tree, feature_names=seeds_features.columns)
```

<!-- grade_id: dt_performance-public, points: 3, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(gini_tree.predict(seeds_df.iloc[:3,0:-1]), [1., 1., 1.])
```

<!-- grade_id: dt_performance-private, points: 1, grade: true, solution: false -->
```python
# Remember there are hidden tests to cross check your implementation. 
```

# 3 Cross Validation and Hyperparmeter Tuning

In this problem you will be using cross-validation to evaluate a model and to tune a hyperparameter.

## 3.1 Example: Basic Cross Validation

Here is an example of using K-fold cross validation for hyperpareter tuning.

```python
from sklearn.model_selection import cross_val_score
from sklearn.model_selection import KFold
```

We next need to define out training dataset and test dataset. Remember we can only use our **training** dataset (`X_train`) to select out hyperparameters, so we'll define that first.

```python
# Initialize a k-fold splitter
kf = KFold(n_splits=3)

# Kf.split() allows you to iterate though the different folds
# "train_index" are the indices of the training data in that fold
# "test_index" are the indices of the testing data in that fold
for train_index, test_index in kf.split(X_train):
    print("Train: ", train_index)
    print("Test: ", test_index)
    print("----")
```

The above code defined 3 folds of data within the training data. We can use those train/validation combinations to  evaluate a given hyperparameter. 

## 3.2 Hyperparameter Tuning with CV (3 pts)

In practice, you generally won't include your test dataset in hyperparameter tuning as you don't want the selection of hyperparameters being associated with the test data. Therefore, in this exercise we'll be trying out different HPs by constructing validation sets from our training data. For conveninece, let's assume test dataset as the validation dataset for the seeds dataset.

<!-- grade_id: k_fold, points: 0, grade: false, solution: true -->
```python
def k_fold_accuracy(k, model, X_data, Y_data):
    """
    You will need to implement a k-fold cross validation with accuracy as the performance metric.
    
    Your inputs and outputs are specified as below:
    
    Input:
        k: An integer indicates the number of fold to do cross validation.
        model: A DecisionTreeClassifier instance.
        X_data: A numpy array of shape (n_data_rows, n_attributes) where n_data_rows refers to 
              the number of rows in your dataset and n_attributes refers to the number of attributes. 
              This typically is from the training dataset.
        Y_data: A numpy array of shape (n_data_rows, ) containing the class labels for each row in your 
              dataset. This typically is from the training dataset.
    Output:
        scores: A list of length k, with scores[i] containing the validation accuracy score for the i-th fold.
    """
    # Init k-fold splitter
    kf = KFold(n_splits=k)
    scores = []
    
    ### BEGIN SOLUTION

    ### END SOLUTION

    return scores
```

```python
# Testing K-fold
k = 3
model = DecisionTreeClassifier(criterion = "gini", random_state=random_seed)
per_fold_acc = k_fold_accuracy(k, model, X_train, Y_train)
print(per_fold_acc)
np.mean(per_fold_acc)
```

<!-- grade_id: k_fold_test, points: 3, grade: true, solution: false -->
```python
# Public tests
k = 3
test_model = DecisionTreeClassifier(criterion = "gini", random_state=234)
per_fold_acc = k_fold_accuracy(k, test_model, X_train, Y_train)
np.testing.assert_almost_equal(per_fold_acc, [0.9464684210526315, 0.9285473684210527, 0.9106973509933775], decimal = 4)
```

<!-- grade_id: k_fold_test_2, points: 1, grade: true, solution: false -->
```python
# Remember there are hiddent tests. 
```

There also exists a built in sklearn function for this, however it is import to know how to perform your own k-fold cross validation split if you want to implement a custom evaluation metric.

```python
from sklearn import metrics
# We're using the trianing dataset here, but remember that CV will
# split that data into training and validation sets for each fold
# so we get an "unbiased" estimate of our test performance.
per_fold_acc = cross_val_score(model, X_train, Y_train, cv=KFold(n_splits=k), scoring='accuracy')
print(per_fold_acc)
np.mean(per_fold_acc)
```

*Why would we ever do this by hand, if there's already a built-in method?* 

Sometimes our model training process is more complex than just fitting the model. For example, we may want to do:
* Feature selection
* Normalization / scaling
* More complex models not in the sklearn library

In these cases, we can still *only use our training data*! You can't use test data to select features - that would be "cheating." So everying that you use your training data for has to occur within the loop we wrote for CV, above, based on the training data for the particular fold we're evaluating.

sklearn also has a feature called [Pipelines](https://scikit-learn.org/stable/modules/generated/sklearn.pipeline.Pipeline.html), which we'll be exploring in the future, to make this easier.

## 3.3 Tuning (10 pts)

In this problem you are going to select the best hypterparameter, using *only the training dataset*. No peaking at the test dataset. To estimate how well a given hyperparameter value will do on *unseen* data, we can use Crossvalidation (within the training dataset) to evaluate our model.

Let's use this approach to select the best `ccp_alpha` hyperparameter for a Decision Tree model.

In the following cell, you should write code to:
1. Iterate over all ccp_alpha values. For each ccp_alpha value, do step 2-3.
2. Calculate the k_fold validation accuracy using the above funciton.
3. Calculate the training accuracy and the validation accuracy.
4. Plot both accuracies vs. the ccp_alpha value.

Don't forget to set your decision tree's **random_state=random_seed** for consistent output.

<!-- grade_id: hyperparameter_search, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import accuracy_score

# np.arange generates a list that starts at minimum, ends at maximum, and increments by step
alpha_values = np.arange(0, 0.035, 0.002)

# two lists to hold our accuracy
k = 5
valid_accs = []
train_accs = []

# TODO: fill the valid_accs and train_accs lists with the model performance for the corresponding ccp_alpha values
### BEGIN SOLUTION

### END SOLUTION    

plt.plot(alpha_values, valid_accs, color='red')
plt.plot(alpha_values, train_accs, color='blue')
plt.xlabel("Post Pruning Alpha")
plt.ylabel(f'Average Accuracy of {k}-fold validation')
```

<!-- grade_id: hyperparameter_search_test, points: 4, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(train_accs, [1.0, 1.0, 0.999865934065934, 0.9882065934065934, 0.9762131868131869, 0.9762131868131869, 0.9762131868131869, 0.9762131868131869, 0.9762285714285714, 0.9762417582417583, 0.9762417582417583, 0.9642417582417583, 0.9642417582417583, 0.9642417582417583, 0.964254945054945, 0.9643703296703297, 0.9643703296703297, 0.9643703296703297], decimal = 4)
np.testing.assert_almost_equal(valid_accs, [0.9105659340659341, 0.9105659340659341, 0.9105659340659341, 0.9105615384615385, 0.9227615384615385, 0.9227593406593406, 0.9288593406593406, 0.9286537362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363, 0.9286237362637363], decimal = 4)
```

**Now answer the following questions:**
1. How does training accuracy change as the alpha value changes? Why? 
2. How does the validation accuracy change as the alpha value chanegs? Why? 
3. Knowing this, what value would you choose for alpha for the final model? Why? Plot the final decision tree with your alpha selection. 

<!-- grade_id: hp_discussion, points: 3, grade: true, solution: true -->
**Discuss here**

The following code selects the alpha value for the best model. Then you job is to train a new model (using all of the training data), using your best hyperparameter value. Then evaluate it on the test dataset using accuracy score. 

```python
# Take the alpha for the model with the best accuracy on the *validation* set!
best_alpha = alpha_values[np.argmax(valid_accs)]
best_alpha
```

<!-- grade_id: train_best_ccp_alpha, points: 0, grade: false, solution: true -->
```python
# Train your decision tree classifier here with ccp_alpha equals to the best alpha find above.
# Your decision tree should use gini index as split critirion.
# Don't forget to set random_state=random_seed.
# You may want to print the tree using plot_tree

gini_tree = None

#TODO: construct the gini_tree
### BEGIN SOLUTION

### END SOLUTION

plt.figure(figsize=(20,10))
_ = plot_tree(gini_tree, feature_names=seeds_features.columns)
```

<!-- grade_id: train_best_ccp_alpha_test_2, points: 1, grade: true, solution: false -->
```python
# Rememer there are hidden tests. 
```

<!-- grade_id: eval_best_ccp_alpha, points: 0, grade: false, solution: true -->
```python
# Now evaluate your model on the test dataset - what are the evaluation metrics?

y_test_accuracy = None
### BEGIN SOLUTION

### END SOLUTION

print(y_test_accuracy)
```

<!-- grade_id: eval_best_ccp_alpha_test, points: 1, grade: true, solution: false -->
```python
assert y_test_accuracy == 0.8571428571428571
```